# Tutorial 10: Indexing, slicing and masking

## PHYS 2600

```python
## Commonly-used import cell

%matplotlib inline

import numpy as np
import matplotlib.pyplot as plt
```

## T10.1 - Indexing, slicing, and masking

Let's begin with some simple practice in indexing.  First, __run the cell below__ to load a few arrays into the global namespace.  (If you make any mistakes below, you can always re-run this cell to re-initialize the arrays.)

```python
planets = np.array(['Mercury', 'Venus', 'Earth', 'Mars', 'Jupiter', 'Saturn', 'Uranus', 'Neptune', 'n/a'])
powers_of_two = np.array([1,   2,   4,   8,  16,  32,  64, 128, 256, 512])
some_data = np.array([ 4.00, 1.72, -0.22, 1.80, -3.61, -0.90, -3.29, 4.45, 3.29, 2.26])
also_planets = planets
```

### Part A

Use _index_ notation `a[i]` to accomplish the following:

* Get and print the name of the sixth planet, Saturn, from `planets`.
* Set the last entry of `planets` to `'Pluto'` (then print `planets` again to make sure it worked.)

<!-- grade_id: cell-d18ef10a6b6141df, points: 2, grade: true, solution: true -->
```python
#
```

When you assigned to `planets` above, you changed the array __in-place__, which means you actually modified the block of memory allocated when we made `planets`.  That means that any name which points to the _same array_ will see the changes you made.

__Print out the array `also_planets`__ and verify that it shows the addition of Pluto as well.

<!-- grade_id: cell-48bc3e3b9c55cb35, points: 0, grade: false, solution: true -->
```python
also_planets
```

### Part B

Use _slice_ notation `a[i:j:k]` to accomplish the following:

* Get and print the names of the _inner planets_, i.e. everything closer to the Sun than Jupiter, from `planets`.
* Get and print all _even_ powers of 2, i.e. $2^0, 2^2, 2^4...$, which appear in `powers_of_two`.


<!-- grade_id: cell-a4a4dbe080e51c1f, points: 0, grade: false, solution: true -->
```python
#
```

Now a more complicated use of slice notation: __divide the last three entries__ in `some_data` __by 5__ (change them __in-place__, i.e. this should _mutate_ `some_data`!)  The test following will verify that you've done this correctly.

(If the `some_data` array gets messed up and you need to reset it, just go back and re-run the first cell in this problem.)

<!-- grade_id: cell-0d7138c61ff51473, points: 0, grade: false, solution: true -->
```python
#
```

<!-- grade_id: cell-499288f8c2a220c9, points: 0, grade: true, solution: false -->
```python
# Testing cell

import numpy.testing as npt

npt.assert_allclose(some_data, [4, 1.72, -0.22, 1.8, -3.61, -0.9, -3.29, 0.89, 0.658, 0.452], rtol=1e-2)
```

### Part C

Finally, use _mask_ notation `a[bool_array]` to accomplish the following:

* Get and print all of the negative entries in `some_data`.
* Square all of the entries in `powers_of_two` which are greater than 100 (then print `powers_of_two`.)

<!-- grade_id: cell-cffc75db17d3134b, points: 0, grade: false, solution: true -->
```python
#
```

## T10.2 - Changing arrays

In the above examples, we made modifications to our arrays _in place_.  Now let's investigate some of the NumPy operations that create _new_ arrays as a result.  As always, use the NumPy documentation and/or the `?` magic command if you're not sure exactly how to use a function.

### Part A

The array `a` below contains a list of even numbers from 0 to 16.  __Use `np.insert()`__ to make a new array with the number 9 added to it (in numerical order.)  Then, print your new array to make sure it worked.

<!-- grade_id: cell-cb560810a7c77c3b, points: 0, grade: true, solution: true -->
```python
a = np.arange(0,17,2)

```

### Part B

__Use `np.delete()`__ to remove Pluto from the array `planets`, defined above.  Then print `planets` again to verify that you have removed it.

(Remember: `np.delete()` returns a new array, just like the other functions in this problem!  You'll have to save the result _back_ to `planets` if you want to change the original.) 

<!-- grade_id: cell-c3cfb2a02999ef43, points: 0, grade: false, solution: true -->
```python
#
```

### Part C

__Use `np.append()`__ to create a single array  that _counts from 0-60_ exactly as follows:

* Every number from 0 to 10 (including 0), then
* Every _other_ number from 10 to 30 (so `[10,12,14...]`), and finally
* Every _third_ number from 30 to 60, including 60 (so `[30,33,36...,60]`).

As a quick way to check your answer, your final array should have length 31.

<!-- grade_id: cell-78fc92a56705f106, points: 0, grade: true, solution: true -->
```python
#
```

## T10.3 - More on masks and Boolean logic

Masks are one of the most powerful tools we have for manipulating parts of a NumPy array, beyond the simple examples we saw above.  Let's get more comfortable with using masks, and then see a couple of applications.

### Part A

Let's begin with the example from lecture:

<img src="https://raw.githubusercontent.com/wlough/CU-Phys2600-Fall2025/main/lectures/img/mask.png" width=400px />

```python
a = np.array([17,-3,-44,8,56])
pos_mask = a > 0
print(pos_mask)
print(a[pos_mask])
```

As we saw in lecture, we can use the mask __in reverse__, to assign to a sub-array of `a`.  In the cell below, __use `pos_mask` to multiply all of the positive entries in `a` by -1__, so that the array becomes entirely negative.

<!-- grade_id: cell-b409d4b737aa9325, points: 0, grade: false, solution: true -->
```python
#
```

When we use a comparison operator like `a > 0`, we _create a new array_ to be a mask.  That means that `pos_mask` doesn't change automatically if we change `a` - __run the cell below__ to verify that `pos_mask` is unchanged.  (Of course, we can always run a new comparison `a > 0` to get an updated mask.)

```python
print(a)
print(pos_mask)
print(a > 0)
```

### Part B

As with single Boolean values, there are many situations where applying logical operations to combine or negate a mask is very useful.  Let's run through some examples: first __run the next cell__ to create an array and some masks.

```python
b = np.arange(21)
div_2_mask = b % 2 == 0
div_3_mask = b % 3 == 0
div_5_mask = b % 5 == 0
print(b)
print(b[div_2_mask])
print(b[div_3_mask])
print(b[div_5_mask])
```

Our masks here are using the modulo operator to single out numbers in `b` that are divisible by 2, 3, or 5.  Thus, for example, `b[div_2_mask]` gives us all the even numbers.  

But what if we want all the _odd_ numbers in `b` instead?  We could make a new mask like `b % 2 != 0`, but we don't need to, since `div_2_mask` already has the information we want if we apply the logical operation `not` to it.

Unfortunately, we can't use the keyword `not`, as you'll see from running the cell below:

```python
not div_2_mask
```

This is the error you'll see if you try to use any of `and`, `or`, or `not` on an array: these operators only work on single Boolean values, not arrays!  Instead, we must use the NumPy equivalent functions.  __Run the next cell__ to see simple examples of each NumPy logical operation.

```python
bool_A = np.array([True, True, False])
bool_B = np.array([False, True, True])

print(np.logical_not(bool_A))
print(np.logical_or(bool_A, bool_B))
print(np.logical_and(bool_A, bool_B))

# Shorthand with "bit logic" operators
print(~bool_A)
print(bool_A | bool_B)
print(bool_A & bool_B)
```

The operators `~`, `|` and `&` are __"bit logic" operations__, that act on the underlying binary data of a Python object directly.  For most purposes you should avoid these operators, but since a Boolean object is a single bit, they just act like shorthand for the logical operations `not`, `or`, `and`.  

You can use them interchangeably, but __be careful with the bit-wise operators__, because they will work on _non-Boolean_ objects too - and the results will be pretty weird!

Back to our example: use the `not` operation and `div_2_mask` to __select all of the odd entries in `b`__.

<!-- grade_id: cell-db07479c86ec9af1, points: 0, grade: false, solution: true -->
```python
#
```

Now use the `and` operation with the masks we defined to __select all of the multiples of 6 in `b`__.

<!-- grade_id: cell-3c6734e7aea2175f, points: 0, grade: false, solution: true -->
```python
#
```

Finally, let's try assigning through the mask: use the `or` operation to __set every entry of `b` which is not a multiple of 2, 3, or 5 to zero.__  (Print `b` again to verify your operation.)

<!-- grade_id: cell-2f7a0e1a9c0d64a7, points: 0, grade: false, solution: true -->
```python
#
```

### Part C

One more bit of weirdness related to NumPy masks before we finish.  We learned before that for convenience, Python will let us chain inequalities together.  This is most useful for testing if a number is inside of a particular interval:

```python
x = 0.5
y = 1.2

print(0 < x < 1, 0 < y < 1)
```

If you think about it, any double-inequality test like this is equivalent to doing the two tests on either side alone, and then joining them with an `and`.  In fact, __that is exactly what Python does__.  Unfortunately, this means that when we try to use this notation with NumPy, we get an error!

__Run the cell below__ to see the familiar message about ambiguous truth values from above.  

```python
z = np.array([0.5, 1.2])
print(0 < z < 1)  ## Interpreted as "(0 < z) and (z < 1)"
```

Now using your new understanding of logic operations in NumPy, __fix the code__: rewrite the test in the cell below so it will give the array `0 < z < 1 = [True, False]` without an error.

<!-- grade_id: cell-2d61efc4a76c54d0, points: 0, grade: false, solution: true -->
```python
#
```

### Part D (optional challenge)

Now let's actually do something useful with masks! Implement the __square wave pulse__ function `square_wave(x)` below, which is defined by

$$
w(x) = \begin{cases} 1, -1 \leq x \leq 1; \\ 0,\ \rm{otherwise}. \end{cases}
$$

_(Hint: the_ `np.zeros_like()` or `np.ones_like()` _functions would be a useful starting point for making the array `w` that your function should return.  Then you just need to apply a mask to set some of the zeros to 1 (or vice-versa.)_

__Important note:__ make sure you make a _new_ array for the result of the function `w(x)`, and don't overwrite anything in `x` itself!  Changing `x` would be an example of a __side effect__, which is when a function does something _other_ than just returning a new object.  If you're trying to plot `w(x)`, you don't want to change the `x` array that you're given!

<!-- grade_id: cell-4c11297537d5960e, points: 0, grade: true, solution: true -->
```python
def square_wave(x):
    #
    
```

<!-- grade_id: cell-de9e8e87f11f2ee5, points: 0, grade: true, solution: false -->
```python
import numpy.testing as npt

x = np.arange(-5,5)
npt.assert_allclose(square_wave(x), [0., 0., 0., 0., 1., 1., 1., 0., 0., 0.])
npt.assert_allclose(x, np.arange(-5,5))  # Make sure x didn't change - no side effects!
```

<!-- grade_id: cell-65e4dfaf2a4f5a01, points: 0, grade: true, solution: false -->
```python
# Test plot: should look like a square pulse from -1 to 1

x = np.linspace(-3,3,100000)
w = square_wave(x)
plt.plot(x,w)
```