# Homework 9

## PHYS 2600

__Important notice:__ All cells in your notebook will be run, start to finish, using a fresh kernel when it is graded! To make sure the graded notebook looks like what you expect, we recommend selecting "Runtime > Restart session and run all" from the menu above in Colab before you finish.

```python
# Import cell
%matplotlib inline

import matplotlib.pyplot as plt
import numpy as np
```

If you worked collaboratively on this assignment, __include the names of your collaborators in the cell below:__

## 9.0 - Homework correction (3 points)

In the cell below, explain and correct __one mistake__ you made on homework #7.  If you got a perfect score, great, tell us what your favorite problem was, or use the space to give any other feedback you might have on the class/tutorials/homework.

<!-- grade_id: hw9-correction, points: 3, grade: true, solution: true -->
```python

```

## 9.1 - The other Newton's law (14 points)

We all know about Newton's three laws of motion, but there's actually another 'law' discovered by and named after Sir Isaac Newton: __Newton's law of cooling__.  The law states that if an object at temperature $T$ exists in an environment with ambient temperature $T_a < T$, then the rate of heat loss can be described using the simple differential equation

$$
\frac{dT}{dt} = -h (T - T_a)
$$

where the (positive) heat loss constant $h$ must be determined experimentally for a given object and environment.

### Part A (4 points)

Write a _discrete_ version of this differential equation using Euler's method, i.e. rewrite it in the form

$$
T_{i+1} = T_i + F(t_i, T_i) \Delta t
$$

<!-- grade_id: hw9_1a_1, points: 2, grade: true, solution: true -->
__Introducing a discrete set of times $t_i$ separated by constant spacing $\Delta t$ and corresponding $T(t_i) = T_i$, the initial differential equation can be rewritten as__

$$
\frac{T_{i+1} - T_i}{\Delta t} = -h(T_i - T_a)
$$

__or rearranging to the desired form,__

$$
T_{i+1} = T_i - h(T_i - T_a) \Delta t
$$

Suppose a freshly brewed mug of coffee has an initial temperature of $T = 85.0^\circ\ C$ (185 degrees F, which is too hot to drink!)  After waiting 30 seconds, you measure the temperature again and it has reduced slightly to $82.5^\circ\ C$.  The room is at a constant temperature of $20^\circ\ C$. 

Using the discrete equation you just derived, __estimate the constant $h$ for the cup of coffee.__

_(Hint: since all you have here are two discrete measurements, you can use the discrete difference equation that you just wrote down to solve for $h$ here.)_

<!-- grade_id: hw9_1a_2, points: 2, grade: true, solution: true -->
__We're trying to find $h$, so we should rearrange the equation again:__

$$
h = \frac{T_i - T_{i+1}}{\Delta t (T_i - T_a)}
$$

__Now we just plug in $\Delta t = 30$ s, $T_i = 85.0$, $T_{i+1} = 82.5$ and $T_a = 20$:__

$$
h = \frac{85.0 - 82.5}{30 (85 - 20)} = \frac{1}{780\ {\rm s}},
$$

__or $h =1.28 \times 10^{-3}\ {\rm s}^{-1}$.__

### Part B (6 points)

__Implement the function `cool_solve` below__, which should carry out a numerical solution of the differential equation you wrote in part A.  (Refer back to lecture/tutorial 13 and the `euler_solve` function we wrote there if you're not sure how to set this up.)

<!-- grade_id: cell-cb3e21c3bae3ffe4, points: 0, grade: false, solution: true -->
```python
def cool_solve(T0, Ta, h, t_max, dt):
    """
    Solve Newton's law of cooling using Euler's method,
    for the given step size dt.

    Arguments:
    ----
    * T0: initial temperature of the cooling object, in degrees C.
    * Ta: ambient temperature of the environment, in degrees C.
    * h: heat constant, in inverse seconds.
    * t_max: how long to run the solution for, in seconds.
    * dt: step size to use, in seconds.

    Returns:
    ----
    * t_sol, T_sol: arrays containing the discretized interval {t_i} and the solution {T_i}.
    """

    #
```

<!-- grade_id: hw9_1b_1, points: 4, grade: true, solution: false -->
```python
# testing cell
import numpy.testing as npt

# Solving for one step with the correct value of h should give back 82.5 as the answer.
t, T = cool_solve(85.0, 20.0, 1 / 780.0, 31.0, 30.0)
print(T)
npt.assert_allclose(T, [85.0, 82.5])

# Solving for a really long time should give about 20 degrees
t, T = cool_solve(85.0, 20.0, 1 / 780.0, 30000.0, 30.0)
print(T[-1])
npt.assert_allclose(T[-1], 20.0, atol=1e-6)
```

Now __make a plot of your solution in the cell below__, using a step size of $dt = 1$ s and evolving for 1 hour (3600 s.)

<!-- grade_id: hw9_1b_2, points: 2, grade: true, solution: true -->
```python
# 
```

### Part C (4 points)

__Answer in the cell below using your code:__  The coffee will be safe to drink at about $60^\circ C$ ($140^\circ F$); how long do you have to wait to have a sip?  How long will it take to reach $25^\circ C$, which is generally considered room temperature?

Again, __use the numerical solution that you plotted in part B (with $dt = 1$ s) to answer this question;__ there is an analytic solution, but I want the Euler's method answer, not the exact analytic answer.  (They're not that far apart, really; you could use the analytic answer to check your work here.)

<!-- grade_id: hw9_1c, points: 4, grade: true, solution: true -->
```python
# 
```

## 9.2 - Vectorized path length (14 points)

Suppose we have an ordered set of points $(x_i, y_i)$ describing a two-dimensional curve.  The corresponding path length of the curve is equal to

$$
L = \sum_{i=1}^n \sqrt{(x_i - x_{i-1})^2 + (y_i - y_{i-1})^2}.
$$

Let's suppose that the set of points is stored as a single $n \times 2$ NumPy array.  ($2 \times n$ might be better for performance, but this will be more human-friendly.)  One way to compute the path length from the array is using a loop, as in the following code:

```python
import math


def path_length_loop(pt_array):
    """
    Calculates the straight-line path length between a list of two-dimensional points
    (x_i, y_i).

    Arguments:
    - pt_array: (n,2) NumPy array containing the points.

    Returns:
    - L: the total path length (float).

    """

    # Keep track of running path length
    L = 0

    # Iterate through all points; we need to work with pairs of points,
    # so use a running index.
    i = 1

    n_pts = pt_array.shape[0]

    # To make the below clearer, make some views showing the columns:
    x = pt_array[:, 0]
    y = pt_array[:, 1]

    while i < n_pts:
        dx = x[i] - x[i - 1]
        dy = y[i] - y[i - 1]

        L += math.sqrt(dx**2 + dy**2)

        i += 1

    return L
```

```python
square_path = np.array(
    [
        [0, 0],
        [1, 0],
        [1, 1],
        [0, 1],
        [0, 0],
    ]
)

print(path_length_loop(square_path))  # Should print "4.0"
```

### Part A  (8 points)

Now it's your turn!  Although the code above will work, using a `while` loop over a NumPy array is actually a _very bad idea!_  This is because looping in Python is very slow compared to just doing vectorized operations over the entire array.

In the cell below, __implement the `path_length_vec` function__, which should function exactly the same as `path_length` above.  However, __your function is _not allowed_ to use any loops (no `while` or `for`!)__  Instead, make use of NumPy functions and operations to compute the path length.

_(Hint: the hard part here is to compute the differences between adjacent x and y points in a vectorized way.  You can have a look at the function `np.diff()`.  Alternatively, you can use the function `np.roll()`, which is a bit more general: it creates a new array where all entries are offset by some number of indices.  For example, `np.roll([0,1,2],1)` gives `[2,0,1]`.  If you use `np.roll` to construct $y_i - y_{i-1}$ as a difference of arrays,  think carefully about what happens at the ends of your array...)_

<!-- grade_id: hw9_2a_1, points: 4, grade: true, solution: true -->
```python
# Remember, no for or while loops allowed in the below function!

import numpy as np


def path_length_vec(pt_array):
    """
    Calculates the straight-line path length between a list of
    two-dimensional points (x_i, y_i).

    Arguments:
    - pt_array: (n,2) NumPy array containing the points.

    Returns:
    - L: the total path length (float).

    """

    # Hint: making views for the x and y points will be helpful again.
    # It is possible to do the job just by working directly with the
    # multi-dimensional pt_array as well!  (This code will be shorter
    # but not really any faster.)

    #
```

<!-- grade_id: hw9_2a_2, points: 4, grade: true, solution: false -->
```python
# Some simple tests, using the already validated path_length function

import numpy.testing as npt

short_path = np.array(
    [
        [0, 0],
        [0, 1],
    ]
)

square_path = np.array(
    [
        [0, 0],
        [1, 0],
        [1, 1],
        [0, 1],
        [0, 0],
    ]
)

tri_path = np.array(
    [
        [0, 0],
        [1, 0],
        [0, 1],
        [0, 0],
    ]
)

npt.assert_allclose(
    path_length_loop(short_path), path_length_vec(short_path)
)  # should be 1
npt.assert_allclose(
    path_length_loop(square_path), path_length_vec(square_path)
)  # should be 4
npt.assert_allclose(
    path_length_loop(tri_path), path_length_vec(tri_path)
)  # should be 2 + sqrt(2)

penta_path = np.array(
    [
        [1.00000000e00, 0.00000000e00],
        [3.09016994e-01, 9.51056516e-01],
        [-8.09016994e-01, 5.87785252e-01],
        [-8.09016994e-01, -5.87785252e-01],
        [3.09016994e-01, -9.51056516e-01],
        [1.00000000e00, 0.000000000e00],
    ]
)

npt.assert_allclose(
    path_length_loop(penta_path), path_length_vec(penta_path)
)  # should be about 5.9
```

### Part B (6 points)

A path-length calculator is simple, but has lots of applications.  One simple one is to calculate the value of $\pi$!  We can inscribe an $n$-sided polygon inside of a circle of radius $r$ by connecting the set of points

$$
x_i = r \cos \left( \frac{2\pi i}{n} \right) \\
y_i = r \sin \left( \frac{2\pi i}{n} \right)
$$

with $i$ running from 0 up to $n$.  As $n$ increases, the path will become very close to a circle and the value of the path length will approach $2\pi$.

This is also an ideal test case to compare our two implementations, since we need a much larger array to really see the performance difference of vectorization!  In the cell below, I've set up a function called `circle_path` which will generate the $n$-gon path in the format we've been using above.  __Complete the function below.__

In the limit that $n$ becomes very large, the path length around your polygon will just be the circumference of the circle.  Use this fact to __estimate $\pi$ from the path length of `circle_path`__; save your answer in the variable `pi_approx`.  _Choose $n$ large enough_ that your estimate is accurate for the __first six digits of $\pi$ = 3.14159...__

<!-- grade_id: cell-040b0fad0dd29a04, points: 0, grade: false, solution: true -->
```python
def circle_path(n, r=1):
    """
    Creates a set of (x,y) points describing
    an n-sided polygon inscribed in a circle of radius r.

    Arguments:
    ---
    n: number of sides the polygon should have.
    r: radius of the circle (default=1.)

    Returns:
    ---
    pt_array: ((n+1)x2) array of coordinates (x_i, y_i),
    where i runs from 0 to n (including both ends.)

    """

    i_array = np.arange(0, n + 1)
    pt_array = np.empty((n + 1, 2))

    #
```

<!-- grade_id: cell-6fe24b4a32064596, points: 0, grade: false, solution: true -->
```python
# Find pi using your path length function, and save it to pi_approx!
pi_approx = 0.0

#
```

<!-- grade_id: hw9_2b_1, points: 3, grade: true, solution: false -->
```python
import numpy.testing as npt

pt_test = circle_path(8)
assert type(pt_test) is np.ndarray
assert pt_test.shape == (9, 2)

penta_path = np.array(
    [
        [1.00000000e00, 0.00000000e00],
        [3.09016994e-01, 9.51056516e-01],
        [-8.09016994e-01, 5.87785252e-01],
        [-8.09016994e-01, -5.87785252e-01],
        [3.09016994e-01, -9.51056516e-01],
        [1.00000000e00, 0.000000000e00],
    ]
)

npt.assert_allclose(circle_path(5, r=1), penta_path, atol=1e-6)
```

<!-- grade_id: hw9_2b_2, points: 3, grade: true, solution: false -->
```python
# testing cell

npt.assert_allclose(path_length_vec(circle_path(4, r=np.sqrt(2))), 8.0)
npt.assert_allclose(pi_approx, np.pi, atol=1e-5)
```

## 9.3 - Radioactive decay (10 points)

If we have a population of particles which undergo radioactive decay, the population over time $N(t)$ will obey the relation

$$
N = N_0 e^{-t/\tau},
$$


where $\tau$ is the lifetime.  This follows from a simple differential equation,
$$
\frac{dN}{dt} = -\frac{N}{\tau}
$$

Actually, this equation has made a hidden assumption: that the number of particles $N$ is so large that writing $dN$ is a sensible thing to do!  The underlying idea has to do with probability: if we observe a _single_ unstable particle for time $dt$, then (from quantum mechanics) the probability we see it decay is

$$
p = 1 - e^{-dt/\tau}.
$$

This is an ideal system to study by Monte Carlo simulation!  Our algorithm will be:

* Start with some population of particles, $N$.
* Take one timestep $dt$, and compute the probability of decay $p$ using $dt$ and $\tau$.
* For each particle, draw a random number $0 < r < 1$.  If $r < p$, then that particle decays.
* Reduce the population $N$ by 1 for each decay that we observe.
* Repeat for the next timestep.

We'll consider the radioactive decay of an isotope of nitrogen, ${}^{16}N$, which has a lifetime of $\tau = 10.286$ s.

### Part A (4 points)

First, __implement the function `decay_step(N, dt)`__ below, which should draw $N$ random numbers and follow the algorithm above for a single timestep.  I've provided some comments to guide you.

Note that when you write the algorithm, the last three steps _could_ be done by a loop; but in practice it's better to draw all $N$ random numbers at once using NumPy, count how many satisfy $r < p$ in the array, and then subtract that number from $N$.  This is the implementation suggested in my comments below.


<!-- grade_id: cell-78bad5402b31ebc0, points: 0, grade: false, solution: true -->
```python
tau_N16 = 10.286


def decay_step(N, dt):
    # Calculate the decay probability p

    # Draw N random numbers between 0 and 1

    # Count the number of decays, by counting
    # how many random numbers in the array are less than p

    # Return N minus the number of decays

    #
```

<!-- grade_id: hw9_3a, points: 4, grade: true, solution: false -->
```python
print(decay_step(100, 0.0001))  # No decay should happen after very short dt

npt.assert_allclose(decay_step(100, 0.0001), 100, atol=1)

N_test = []
for _ in range(100):
    N_test.append(decay_step(100, tau_N16))

N_tau_mean = np.mean(N_test)

# Result should always be within 6-sigma of standard error
# Standard error is about 0.5
# Odds of random failure ~ 1 per billion
print(N_tau_mean)
npt.assert_allclose(N_tau_mean, 37, atol=3)
```

### Part B (6 points)

Next, __implement the function `monte_carlo_decay(N0, t)`__, which takes the initial number of particles `N0` and an array `t_array` of discrete times to track the particles over.  Your function should return an array `N_vs_t` of the same length as `t_array`, containing the number of particles remaining at each discrete time.

_(Hint: you may want to look back at tutorial 13 again; the way that `monte_carlo_decay()` should work with `decay_step()` is similar to how `euler_solve()` built up the full solution to the ODE on that tutorial using `euler_step()`.)_

<!-- grade_id: cell-114c4c0b8d52b5f2, points: 0, grade: false, solution: true -->
```python
def monte_carlo_decay(N0, t_array):

    # Get dt (step size) and Nt (number of t points) from t_array
    dt, Nt = t_array[1] - t_array[0], len(t_array)

    # Initialize array N_t, same length as t_array
    # The "dtype" option is important, because the number
    # of particles should always be an integer!
    N_vs_t = np.zeros_like(t_array, dtype=int)

    # Set the initial (first) value of N_t to value N0 passed in above

    # Iterate, setting N_t[i+1] using the decay_step function and N_t[i]

    #

    ## Return array N(t)
    return N_vs_t
```

Now __run the cell below__ to plot your simulation results for an initial population 1000 nitrogen-16 atoms over 60 seconds.  I've also included a plot of the analytic result for $N(t)$ as a red, dashed line.  Hopefully the two curves agree pretty closely!

<!-- grade_id: hw9_3b_1, points: 2, grade: true, solution: true -->
```python
t_decay = np.arange(0, 60, 0.01)
N0 = 1000

N_analytic = N0 * np.exp(-t_decay / tau_N16)
plt.plot(t_decay, N_analytic, linestyle="--", color="r")

N_N16 = monte_carlo_decay(N0, t_decay)
plt.plot(t_decay, N_N16)
```

<!-- grade_id: hw9_3b_2, points: 2, grade: true, solution: false -->
```python
# Verify that the evolution is very close to analytic solution while N is relatively large

npt.assert_allclose(N_N16[:1000], N_analytic[:1000], rtol=2e-1)
```

Finally, in the cell below __copy the code from the plotting cell above__, change `N0` to 5, and then run it again to show your simulation vs. the analytic formula.  In the answer cell below, explain: which of the two models do you think is a more realistic description of the system starting with 5 particles, and why?

<!-- grade_id: cell-31d40aa5fe1625d5, points: 0, grade: false, solution: true -->
```python
# 

# The discrete description is more realistic, since we really only have 5
# particles, so N can't be fractional like the smooth analytic curve predicts!
```

<!-- grade_id: hw9_3b_3, points: 2, grade: true, solution: true -->
__The discrete curve (blue) is more realistic!  We really only have 5 particles, so the population $N$ can't actually be fractional like the analytic curve (red) shows.__