# Homework 5

## PHYS 2600

__Important notice:__ All cells in your notebook will be run, start to finish, using a fresh kernel when it is graded! To make sure the graded notebook looks like what you expect, we recommend selecting "Runtime > Restart session and run all" from the menu above in Colab before you finish.

```python
## Import cell

%matplotlib inline
import numpy as np
```

If you worked collaboratively on this assignment, __include the names of your collaborators in the cell below:__

## 5.0 - Homework correction (3 points)

In the cell below, explain and correct __one mistake__ you made on homework #3.  If you got a perfect score, great, tell us what your favorite problem was, or use the space to give any other feedback you might have on the class/tutorials/homework.

<!-- grade_id: hw6-correction, points: 3, grade: true, solution: true -->


## 5.1 - More Boolean logic (9 points)

For the problems below, you should __write a Boolean test function__, i.e. a function that returns a Boolean value `True` or `False`.
 
### Part A (4 points)

Implement the function `one_is_odd(x,y)`, which should return `True` if just one of `x` and `y` is an odd number (i.e. not divisible by 2), and `False` otherwise.

<!-- grade_id: cell-a5c0ff0849f24f75, points: 0, grade: false, solution: true -->
```python
def one_is_odd(x,y):
    ### BEGIN SOLUTION
    return (x % 2 != 0) != (y % 2 != 0)
    ### END SOLUTION
```

Some of the tests below are written incorrectly. Fix them so that they will run without error on your (correctly written!) `one_is_odd` function.

<!-- grade_id: cell-ef92213313962c89, points: 3, grade: true, solution: false -->
```python
assert one_is_odd(2,3)
assert not one_is_odd(5,7)
assert not one_is_odd (2,4)
assert not one_is_odd(6,18)
```

### Part B (4 points)

Implement the function `close_to_sqrt(x,y,tol)`, which should check whether either number is the square root of the other number.  (Specifically, it should return `True` if $x^2$ is within `tol` of $y$ or if $y^2$ is within `tol` if $x$.)

<!-- grade_id: cell-b0c5c0a3b12e63cd, points: 0, grade: false, solution: true -->
```python
def close_to_sqrt(x,y,tol=1e-2):
    ### BEGIN SOLUTION
    return abs(x**2 - y) < tol or abs(y**2 - x) < tol
    ### END SOLUTION
```

<!-- grade_id: cell-45d83211ab7d7f7e, points: 3, grade: true, solution: false -->
```python
assert close_to_sqrt(5.0, 25.0)
assert not close_to_sqrt(2.0, 9.0)
assert close_to_sqrt(36.0, 6.0)
assert close_to_sqrt(49.0, -7.0001, tol=1e-2)
```

### Part C (4 points)

Implement the function `are_negative_ordered(a,b,c)`, which should check to see whether __both__ $a<b<c$ __and__ all three numbers are negative.  (You have to write the function header this time!)

<!-- grade_id: cell-85d296c88a662174, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def are_negative_ordered(a,b,c):
    # Check ordering
    are_ordered = a < b < c

    # Check if c is negative
    c_negative = c < 0

    # If both conditions are true, then b and c are also positive, so we're done!
    return c_negative and are_ordered
### END SOLUTION

```

<!-- grade_id: cell-038e80cd56f7f8f4, points: 3, grade: true, solution: false -->
```python
assert are_negative_ordered(-3,-2,-1)
assert not are_negative_ordered(-4,-3,-5)
assert not are_negative_ordered(-7,3,7)
assert are_negative_ordered(-9,-5,-3)
assert not are_negative_ordered(0,1,1)
assert not are_negative_ordered(-1,-10,-100)
```

## 5.2 - Working with integrals (9 points)

For this problem you'll be evaluating some integrals numerically, using the `np.trapezoid` function.  Choosing the number of points for your discrete grid is up to you - but your answers _should be accurate to within 0.1%_ of the exact result.  __Check that your answer is accurate, and use more points in your grid if it's not, or you will lose points!!__

_(Hint: how can you check your answer?  Remember, a simple test is just to increase the number of points on your grid, and see if the answer changes or not.  This generally works even for cases where you don't know the right answer.)_

_(Another hint: the first integral is easy to work out by hand, so you know the right answer...the second integral also has an analytic answer, although it's harder to get.)_ 

### Part A (3 points)

Evaluate the integral:

$$
I = \int_0^4 \left[ 4(x^3 - 1) \right] dx
$$

__Save your numerical answer to the variable `integral_I`.__  The visible testing cell will do a simple check that your integral is positive, it's up to you to verify the answer...

<!-- grade_id: cell-cc9e5975e7cd6838, points: 0, grade: false, solution: true -->
```python
integral_I = 0.

### BEGIN SOLUTION
xI = np.linspace(0,4,50)
integral_I = np.trapezoid(4*xI**3 - 4, x=xI)
### END SOLUTION
```

<!-- grade_id: cell-4f37e09dbc6bd44f, points: 1, grade: true, solution: false -->
```python
# Basic sanity check - not a thorough test!
print(integral_I)
assert integral_I > 0.
```

Output:
```
240.106622240733
```

<!-- grade_id: cell-0d1ceefdc0e0d996, points: 2, grade: true, solution: false -->
```python
# Hidden testing cell - do not delete!
import numpy.testing as npt

### BEGIN HIDDEN TESTS
npt.assert_allclose(integral_I, 240.0, rtol=0.001)
### END HIDDEN TESTS
```


### Part B (3 points)

Evaluate the integral:

$$
J = \int_0^{2\pi} \sin(2x) e^{-x/2} dx
$$

__Save your numerical answer to the variable `integral_J`.__  The visible testing cell will do a simple check that your integral is positive, it's up to you to verify the answer...

<!-- grade_id: cell-94ecc32dd20c782b, points: 0, grade: false, solution: true -->
```python
integral_J = 0.

### BEGIN SOLUTION
xJ = np.linspace(0,2*np.pi, 200)
integral_J = np.trapezoid(np.sin(2*xJ) * np.exp(-xJ/2), x=xJ)
### END SOLUTION
```

<!-- grade_id: cell-b5ff79ce5bc71b63, points: 1, grade: true, solution: false -->
```python
# Basic sanity check - not a thorough test!
print(integral_J)
assert integral_J > 0.
```

Output:
```
0.4500932944384325
```

<!-- grade_id: cell-d2cb69f30cd2e018, points: 2, grade: true, solution: false -->
```python
# Hidden testing cell - do not delete!

### BEGIN HIDDEN TESTS
npt.assert_allclose(integral_J, (8/17)*(1-np.exp(-np.pi) ), rtol=0.001)
### END HIDDEN TESTS
```

### Part C (3 points)

Evaluate the integral:

$$
K = \int_{-1}^{1} \frac{\sin(500x)}{x} dx
$$

__Save your numerical answer to the variable `integral_K`.__  The visible testing cell will do a simple check that your integral is positive, it's up to you to verify the answer...

_(Hint: integral K is oscillating back and forth pretty badly - you need to be especially careful that you used enough points here!)_

<!-- grade_id: cell-08919c8cfff5dc63, points: 0, grade: false, solution: true -->
```python
integral_K = 0.

### BEGIN SOLUTION
xK = np.linspace(-1, 1, 500)
integral_K = np.trapezoid(np.sin(500*xK) / xK, x=xK)
### END SOLUTION
```

<!-- grade_id: cell-0c593d22c9da203c, points: 1, grade: true, solution: false -->
```python
# Basic sanity check - not a thorough test!
print(integral_K)
assert integral_K > 0.
```

Output:
```
3.143862508895048
```

<!-- grade_id: cell-4da8ccfb43d952c8, points: 2, grade: true, solution: false -->
```python
# Hidden testing cell - do not delete!
import numpy.testing as npt

### BEGIN HIDDEN TESTS
npt.assert_allclose(integral_K, 3.145, rtol=0.001)
### END HIDDEN TESTS
```

## 5.3 - NumPy and precision (6 points)

Let's use NumPy to explore issues related to floating-point precision some more.

### Part A (2 points)

In the cell below, __use NumPy to calculate__ $\sin(0)$ and $\sin(\pi)$, and print them both out.


<!-- grade_id: cell-b3c5efc0587e8624, points: 0, grade: false, solution: true -->
```python
import numpy as np

print(np.sin(0))
print(np.sin(np.pi))
```

Output:
```
0.0
1.2246467991473532e-16
```

Now based on what you know about floating-point precision, explain __why these two values are different__ - and __explain the size of the difference__.

<!-- grade_id: cell-004c52cde60ab9a6, points: 2, grade: true, solution: true -->
__Zero is an exact number, even in floating point, so sin(0) is exactly 0 too.  However, $\pi$ is not exact - it must be truncated as a floating-point number.  With 64-bit precision, we expect the truncation error to be roughly around $10^{-16}$, exactly what we find.__

### Part B (2 points)

Although the truncation error we saw above was very small, it can still be significant if we're not careful.  In particular, if a truncation error occurs many times, we can end up with __accumulated errors__ that can be significant.

In the cell below, __create an array containing the value $\pi$ 100 million times__.  Then take the sine of the array, and sum over it with `np.sum()`.  Print the result.

<!-- grade_id: cell-6765b76ea3ef7964, points: 1, grade: true, solution: true -->
```python
long_array = np.ones(100000000) * np.pi
np.sum(np.sin(long_array))
```

Output:
```
np.float64(1.2246467991473531e-08)
```

Once again, the sum _should_ be zero.  __Explain the size of the difference.__

<!-- grade_id: cell-b9e7c77f9ebfedee, points: 1, grade: true, solution: true -->
__This is a simple extension of what we saw in part A; the truncation error for $\sin(\pi)$ is $10^{-16}$, so if we repeat that error 100 million = $10^{8}$ times, the resulting error is about $10^{-8}$, matching the print-out above.__

### Part C (3 points)

We're insulated from these issues with precision by using 64-bit floating point, but NumPy supports other representations that use small amounts of memory.  

Use the `np.array()` function to create another array with 100 million copies of $\pi$, but this time __in 32-bit precision (using data type `np.float32`.)__  You should get a much larger accumulated error!

<!-- grade_id: cell-5b23f603432911f3, points: 1, grade: true, solution: true -->
```python
short_array = np.array(long_array, np.float32)
np.sum(np.sin(short_array))

```

Output:
```
np.float32(-8.742277)
```

Now __explain the size of the error__ that you find with 32-bit precision!  

_(Hint: look back at lecture 8 to see where the truncation error of approximately $10^{-16}$ came from for 64-bit floats.  A 32-bit float has up to 23 bits of precision in its significant digits.)_

<!-- grade_id: cell-a943363ae2743601, points: 2, grade: true, solution: true -->
__Using the information given above, the maximum truncation error with 32 bits is $(1/2)^{23} \approx 6 \times 10^{-8}$.  If we multiply this by 100 million, we get 6 as the max expected error - which is pretty close to what we see in the output above!__

## 5.4 - Intro to TeX math notation (12 points)

All the way back in lecture 1, I told you that we'd be meeting not just Markdown, but a way to write math equations inside Jupyter Markdown blocks called __MathJax__.  Well, now is the time!

MathJax is a tool that implements the widely-used __TeX math notation__.  "TeX" (pronounced like "tech") is a typesetting language used to make PDF files: it's the gold standard for writing scientific papers.  In its full glory, TeX is a complete markup langauge which is significantly more complicated than Markdown.  MathJax just borrows the math notation, and lets us use it alongside Markdown - the best of both worlds!

You can even use it in things like legends and labels for PyPlot, as in the cell above!  For PyPlot, _an extra `\` appears in all TeX commands that use one_  (this is due to `\` being an "escape character" for PyPlot itself.)

Having both Markdown and TeX math available lets you make complete, self-contained Jupyter notebooks for projects - like every lecture and assignment you've seen in the class so far.

This cell shows the notation for MathJax inside Jupyter.  We can write both __in-line math__, which appears as part of some other text, or __stand-alone equations__ which are separated out.  In-line math looks like this: $a^2 + b^2 = c^2$, and is denoted by surrounding our math with the dollar `$` character.  Stand-alone equations look like this:

$$
\int_0^1 x = \frac{1}{2}
$$

and are denoted by a _pair_ of dollar signs, `$$`.
__Double-click on this cell to see what's inside!__

The _next_ couple cells show some common TeX commands and how they appear.  There's a [nice and somewhat more complete list on StackExchange here](https://math.meta.stackexchange.com/questions/5020/mathjax-basic-tutorial-and-quick-reference).

Basically:
* Regular letters and math symbols like `+`, `-`, `/` are all drawn nicely.
* TeX commands are denoted by a backslash, i.e. as `\command`.  These draw special symbols like $\zeta$, names of special functions like $\sin$, and things like integrals, fractions, sums...
* TeX uses curly braces `{`,`}` to group things together: this tells it the difference between $10^10$ and $10^{10}$.

Even I don't know all of the TeX special commands!  If you're looking for the TeX to draw something specific, [DeTeXify](http://detexify.kirelabs.org/classify.html) will let you draw it and try to guess what command you're looking for!

| Symbol | TeX command | Example |
|--------|-------------|---------|
| superscript  |  `^`        |   $x^2$   |
| subscript |  `_`      |  $x_i$  |
| symbols  | `\alpha, \beta, \Gamma...` | $\alpha$, $\beta$, $\Gamma$ |
| more symbols| `\infty, \lt, \gt, \cdot , \rightarrow...` | $\infty, \lt, \gt, \cdot, \rightarrow$ |
| special funcs | `\sin, \log, \lim...` | $\sin(x)$, $\log(y)$, $\lim_{x \rightarrow 0}$ |
| groups   | `{...}`    |  $10^10$ (wrong) vs. $10^{10}$ (grouped!) |
| fraction | `\frac{a}{b}` |  $\frac{3}{4}$  |
| integral | `\int_{a}^{b}` |  $\int_{-1}^{1}$ |
| derivative | `\frac{dy}{dx}` | $\frac{dy}{dx}$ |
| sum  | `\sum_{a}^{b}`    | $\sum_{x=1}^{10}$ |
| product  | `\prod_{a}^{b}`    | $\prod_{x=1}^{10}$ |
| radicals  | `\sqrt{...}`  | $\sqrt{3}$ or $\sqrt[4]{4}$ |
| accents and fonts | `\vec{x}, \mathbf{y}, \hat{z}...` | $\vec{x}, \mathbf{y}, \hat{z}$ |
| text inside math | `\text{some text}` | $x_{\text{long text label}}$ |


**Brackets/parenthesis:**

| Type | TeX command | Example |
|--------|-------------|---------|
| round 	  |  `(x+y)`        |   $(x+y)$   |
| square 	  |  `[x+y]`        |   $[x+y]$   |
| curly 	  |  `\{ x+y \}`        |   $\{x+y\}$   |
| angled 	  |  `\langle x+y \rangle`        |   $\langle x+y \rangle$   |
| vertical 	  |  `\|x+y\|`        |   $|x+y|$   |

The `\left` and `\right` commands can be used to automatically scale/adjust the size of brackets. Use these when the contents are tall (fractions, sums, matrices). 

| TeX command | Example |
|-------------|---------|
|  `(\frac{1}{2})`        |   $\displaystyle(\frac{1}{2})$   |
|  `\left(\frac{1}{2}\right)`        |   $\displaystyle\left(\frac{1}{2}\right)$   |
|  `\|\frac{1}{2}\rangle`        |   $\displaystyle|\frac{1}{2}\rangle$   |
|  `\left\|\frac{1}{2}\right\rangle`        |   $\displaystyle\left|\frac{1}{2}\right\rangle$   |


### Part A (8 points)

__Write all of the following formulas out in TeX in the cell below:__ use the variable names as described!

* __The Pythagorean theorem:__ the sum of squares of two sides $a$ and $b$ of a right triangle equals the square of the length of the hypotenuse $c$.
* __Newton's second law:__ in one dimension, force vector $F$ is equal to the  derivative of momentum vector $p$ with respect to time $t$.
* __Simple integration:__ the integral of the function 1/x from 1 to 2 is equal to the logarithm of 2.
* __The quadratic formula:__ solutions to the quadratic equation $ax^2 + bx + c = 0$ are minus b, plus or minus the square root of b squared minus 4 times a times c, all divided by 2 times a.

<!-- grade_id: cell-19afc2343368c191, points: 8, grade: true, solution: true -->
* __Pythagorean theorem:__

$$
a^2 + b^2 = c^2
$$

* __Newton's second law:__

$$
\vec{F} = \frac{d \vec{p}}{dt}
$$

* __A simple integral:__

$$
\int_1^2 dx\ \frac{1}{x} = \log(2)
$$

* __The quadratic formula:__

$$
ax^2 + bx + c = 0 \Rightarrow x = \frac{-b \pm \sqrt{b^2 - 4ac}}{2a}
$$

__(You only had to write the formula itself, but I thought I'd show writing both and using `\Rightarrow`.)__

### Part B (4 points)

Now let's put what you know together: in the final cell below, __pick your favorite physics or math equation__ (not on the list you just did), write it out in TeX notation, and then _describe it briefly using Markdown._

<!-- grade_id: cell-09918bb4842399b3, points: 4, grade: true, solution: true -->
__I like the [Baker-Campbell-Hausdorff formula](https://en.wikipedia.org/wiki/Baker%E2%80%93Campbell%E2%80%93Hausdorff_formula), myself:__

$$
\exp(X)\exp(Y) = \exp\left(X + Y + \frac{1}{2} \left[X,Y\right] + \frac{1}{12} \left[X,\left[X,Y\right]\right] + \frac{1}{12} \left[Y,\left[Y,X\right]\right] + \cdots\right)
$$


