# Homework #3

## PHYS 2600

__Important notice:__ All cells in your notebook will be run, start to finish, using a fresh kernel when it is graded! To make sure the graded notebook looks like what you expect, we recommend selecting "Runtime > Restart session and run all" from the menu above in Colab before you finish.

If you worked collaboratively on this assignment, __include the names of your collaborators in the cell below:__

<!-- grade_id: cell-6aad98798978a240, points: 0, grade: true, solution: true -->


## 3.0 - Homework correction (3 points)

In the cell below, explain and correct __one mistake__ you made on homework #1.  If you got a perfect score, great!  Tell us what your favorite problem was, or use the space to give any other feedback you might have on the class/tutorials/homework.

<!-- grade_id: cell-1689332322e959be, points: 3, grade: true, solution: true -->


## 3.1 - Some exercises with NumPy (13 points)

The following are small problems in working with arrays using the `numpy` module.

### Part A (2 points)

__Create an array `cos_squared`__, which contains the value of the function $\cos^2(t)$, for $t$ values ranging from 0 to 6 (including both ends!) and spaced by 0.5.  (So `t = [0, 0.5, 1.0, ..., 6.0].`)

<!-- grade_id: cell-cf85e3b4b07f39d3, points: 0, grade: false, solution: true -->
```python
import numpy as np

### BEGIN SOLUTION
t = np.arange(0, 6.5, 0.5)
cos_squared = np.cos(t)**2
### END SOLUTION
```

<!-- grade_id: cell-ca26a7b2f90181e1, points: 2, grade: true, solution: false -->
```python
import numpy.testing as npt

print(cos_squared)

assert len(cos_squared) == 13
assert type(cos_squared) is np.ndarray

### BEGIN HIDDEN TESTS
cos2_correct = np.array([1.0,         0.77015115, 0.29192658, 0.00500375, 0.17317819, 0.64183109,
 0.98008514, 0.87695113, 0.42724998, 0.04443487, 0.08046424, 0.50221285,
 0.92192698])

npt.assert_allclose(cos_squared, cos2_correct, atol=1e-6)
### END HIDDEN TESTS
```

### Part B (2 points)

__Create an array `length_40_array`__ which, as the name implies, should have 40 entries: the first 20 entries should be the number 2, and the second 20 entries should be the number 20.

_(Hint: there are quite a lot of ways to do this!  The `np.concatenate` function might be worth a look...)_

<!-- grade_id: cell-39bf7e4f98f51b4d, points: 0, grade: false, solution: true -->
```python
import numpy as np

### BEGIN SOLUTION
two_array = np.ones(20) * 2
twenty_array = np.ones(20) * 20
length_40_array = np.concatenate((two_array, twenty_array))
### END SOLUTION
```

<!-- grade_id: cell-6d6de2fd1b899a06, points: 2, grade: true, solution: false -->
```python
import numpy.testing as npt

print(length_40_array)

assert len(length_40_array) == 40
assert type(length_40_array) is np.ndarray

### BEGIN HIDDEN TESTS
npt.assert_allclose(length_40_array[:20], np.ones(20) * 2)
npt.assert_allclose(length_40_array[20:], np.ones(20) * 20)
### END HIDDEN TESTS
```

### Part C (3 points)

A _parallelepiped_ is a three-dimensional shape which is a cousin of the cube, except that its three pairs of identical faces are parallelograms - see [the Wikipedia page](https://en.wikipedia.org/wiki/Parallelepiped) for a sketch and in-depth explanation.  If we describe a parallelepiped in terms of the three vectors $\mathbf{a}, \mathbf{b}, \mathbf{c}$ describing its edges at one corner, then its volume can be computed from the _triple product_

$$
V = |\mathbf{a} \cdot (\mathbf{b} \times \mathbf{c})|.
$$

__Implement the function `triple_product(a,b,c)` below__ using `numpy`, and then run the test cell to calculate the volume of some parallelepipeds described by the given vectors.

<!-- grade_id: cell-d4113d75a420a9bc, points: 0, grade: false, solution: true -->
```python
import numpy as np

def triple_product(a,b,c):
    ### BEGIN SOLUTION
    cross = np.cross(b,c)
    triple = np.dot(a,cross)

    return np.abs(triple)
    ### END SOLUTION
```

<!-- grade_id: cell-69ec9a15411a9eba, points: 3, grade: true, solution: false -->
```python
import numpy.testing as npt

# First test: a cube!  Volume should be 1.
a = [1,0,0]
b = [0,1,0]
c = [0,0,1]

print(triple_product(a,b,c))
npt.assert_allclose(triple_product(a,b,c), 1)

# Second test: a rectangular prism!  Volume should be 4.
a = [2,0,0]
b = [0,2,0]
c = [0,0,1]

print(triple_product(a,b,c))
npt.assert_allclose(triple_product(a,b,c), 4)

# Third test: all angles are 45 degrees, a "rhombohedron" with unit side length
# Volume should be 0.455090, to six digits - used formula from Wikipedia.
a = [1,0,0]
b = [0.70710678, 0.70710678, 0.        ]
c = [0.70710678, 0.29289322, 0.64359425]

print(triple_product(a,b,c))
npt.assert_allclose(triple_product(a,b,c), 0.455090, atol=1e-6)
```

### Part D (3 points)

Now a bit of "data analysis".  The array `photo_stop_volts` below contains a set of measurements of the stopping voltage $V_{\rm stop}$ (in volts) for a few different materials in a photoelectric effect experiment. (Note that __the stopping voltage is always negative__ in our convention.)  The experiment used ultraviolet light, with a frequency $f$ of 1590 terahertz (THz).  

Your task: __create a new array `photo_work_functions`__ that contains the work function $\Phi$ (__in electron-volts, eV__) for each material.

To save you some trouble, here's the relevant equation for the photoelectric effect - you're on your own to look up any necessary physical constants and unit conversions.

$$
\Phi = hf + eV_{\rm stop}
$$

Here $h$ is Planck's constant (not the reduced constant $\hbar$!) and $e$ is the charge of a single electron.

_(Hint: switching to electron-volts instead of SI units as fast as possible will make your life easier...)_



<!-- grade_id: cell-61bd89782dd6f99a, points: 0, grade: false, solution: true -->
```python
import numpy as np

photo_stop_volts = np.array([-3.12, -2.46, -1.62, -3.23])

### BEGIN SOLUTION
# f in THz --> convert to Hz --> multiply by h in eV * s
hf = 1590 * 1e12 * 4.136 * 1e-15
print(hf)

# Work function using the formula above.
# No conversion needed: photo_stop_volts * e is already in electron-volts!
photo_work_functions = photo_stop_volts + hf
print(photo_work_functions)
### END SOLUTION
```

<!-- grade_id: cell-1ae16d474e10d89e, points: 3, grade: true, solution: false -->
```python
import numpy.testing as npt

print("Work functions: ", photo_work_functions)

assert len(photo_work_functions) == 4
assert type(photo_work_functions) is np.ndarray

### BEGIN HIDDEN TESTS
npt.assert_allclose(photo_work_functions, np.array([3.46, 4.12, 4.96, 3.35]), atol=1e-2)
### END HIDDEN TESTS
```

### Part E (3 points)

When working with real data, it is often useful to reduce a long list of individual measurements to a set of _summary statistics_.  The most commonly-used statistics are the mean and the standard deviation; since these are so common, they are available in NumPy.

In the cell below, __implement the function `mean_std_error(x)`__, which should return a NumPy array of length three containing the mean, the standard deviation, and the standard error computed from the input array `x`.  (As a reminder, the _standard error_ of a set of samples is equal to the standard deviation $\sigma$ divided by the square root of the number of samples $N$)

$$
\sigma_{SE} = \frac{\sigma}{\sqrt{N}}
$$

Here, $N$ is the length of the array $x$.  __Make sure__ that you use the usual [bias-corrected estimator](https://en.wikipedia.org/wiki/Standard_deviation#Corrected_sample_standard_deviation) for the standard deviation, which is divided by $\sqrt{N-1}$ and not just $\sqrt{N}$ - you may have to set an optional argument to do this.

_(Hint: to make the required array, first calculate the mean and the standard error and save them to variables.  There are NumPy functions for the mean and the standard deviation that you can use - look them up!  Then return a new `np.array([])` with the mean and the error that you calculated.)_


<!-- grade_id: cell-20eb62e35471b121, points: 0, grade: false, solution: true -->
```python
def mean_std_error(x):
    ### BEGIN SOLUTION
    N = len(x)

    mu = np.mean(x)
    s = np.std(x, ddof=1)

    return np.array([mu, s, s/np.sqrt(N)])
    ### END SOLUTION

```

<!-- grade_id: cell-10852649131dbd5c, points: 3, grade: true, solution: false -->
```python
import numpy.testing as npt

x1_test = np.array([1, 0, -1])
npt.assert_allclose(mean_std_error(x1_test), np.array([0.0, 1.0, 0.577]), atol=1e-2)

### BEGIN HIDDEN TESTS
x2_test = np.array([-3, -1, 0, 1, 2])
npt.assert_allclose(mean_std_error(x2_test), np.array([-0.20, 1.92, 0.86]), atol=1e-2)
### END HIDDEN TESTS
```

## 3.2 - Lots of plots (16 points)

### Part A (3 points)

Plot the function 

$$
f(x) = e^{x/2} + \cos(2x) - x^3/10
$$ 

over the range $-2 \leq x \leq 2$.  __On the same plot__, include the function 

$$
g(x) = e^{x/2}.
$$

__Your plot should satisfy the following:__

- Color: $f$ should be blue, and $g$ should be red.
- Markers: no markers shown, just the lines.
- Style: $f$ should be a solid line, and $g$ should be a dotted or dashed line.
- Include a legend which shows clearly which function corresponds to which curve!

<!-- grade_id: cell-6b76e882983a7ec2, points: 3, grade: true, solution: true -->
```python
import numpy as np
import matplotlib.pyplot as plt

### BEGIN SOLUTION
x = np.linspace(-2,2,50)
f = np.exp(x/2) + np.cos(2*x) + 0.1 * x**3
g = np.exp(x/2)

plt.plot(x,f,color='blue', label='f(x)')
plt.plot(x,g,color='red',linestyle='-.', label='g(x)')
plt.legend()
### END SOLUTION
```

### Part B (3 points)

Since `plt.plot()` simply takes lists of $x$ and $y$ coordinates for points, we've seen that it can easily deal with either lists of data, or functions plotted using `np.linspace`.  In fact, we can also easily use it with __parametric plots__, where both $x$ and $y$ are functions of some auxiliary variable $t$.

__Make a plot of an ellipse below__ using `plt.plot`.  The equation for the ellipse you should plot is:

$$
x = 1 + 2 \cos(t) \\
y = 3 \sin(t)/2
$$

You can choose whatever range of $t$ you want, but make sure your ellipse doesn't have any gaps or jagged edges!

To make it obvious that your ellipse is, in fact, an ellipse, __show your plot with both x and y ranging from -3 to 3.__ Show and label the axes to make the range clear.

<!-- grade_id: cell-b684e3a310359762, points: 3, grade: true, solution: true -->
```python
import numpy as np
import matplotlib.pyplot as plt

### BEGIN SOLUTION
t = np.linspace(0, 2*np.pi, 50)
ellipse_x = 1 + 2*np.cos(t)
ellipse_y = 1.5 * np.sin(t)

plt.plot(ellipse_x, ellipse_y, color='orange')
# Draw axes to make it look a little nicer
plt.axhline(0)
plt.axvline(0)
plt.xlim(-3,3)
plt.ylim(-3,3)
### END SOLUTION
```

### Part C (5 points)

A certain type of _wave packet_ is defined by the function

$$
y(x,t) = e^{-(x-3t)^2} \sin(3\pi (x-3t))
$$

This describes a "traveling wave" propagating through some medium; the shape of the packet remains fixed as it travels.  To see this, let's make a plot!  __Your plot should have the following content:__

* It should be a plot versus $x$, over the range $-10 \leq x \leq 10$.
* It should show the "wave packet" function $y(x,t)$ above versus $x$, for __five values of t__: -2, -1, 0, 1, and 2.
* Each plot $y(x,t)$ for a given value of $t$ should be drawn as a different color.  (Pyplot should do this for you by default!)
* The plot should have a legend of the form $t=X$, labeling each curve with the value of $t$ _as an integer_.  (Remember string formatting?  Now we really need it to get a nice-looking plot legend!)

To discourage you from just using copy-and-paste, I've included a function stub `plot_wave_packet(x,t)` below, which you can implement to create the plot.  In fact, __you're only allowed to call `plt.plot()` inside of `plot_wave_packet(x,t)`__, just to make sure you're not tempted to copy-and-paste.  (To iterate over the five values of `t` with this function, a `for` loop would be helpful.)

_(Note on functions and `pyplot`: since `pyplot` in general uses a single "global state" per cell, __your function doesn't have to return anything at all__; it can just operate on the global state with `plt.plot()` and other operations, and the plot will be created and shown in whatever cell you call_ `plot_wave_packet` _from.)_

If you're bothered by how the legend most likely ends up overlapping part of your plot, have a look at the `bbox_to_anchor=...` keyword argument to `plt.legend()`, and see if you can fix the problem.  (This part is optional, but a useful thing to know about!)

<!-- grade_id: cell-ff96f1c873b02983, points: 5, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt
import numpy as np

def plot_wave_packet(x, t):
    ### BEGIN SOLUTION
    u = (x-3*t)
    y = np.exp(-u**2) * np.sin(3*np.pi*(x-3*t))

    plt.plot(x,y, label='t={:d}'.format(t))
    plt.legend(bbox_to_anchor=(1.3,1.1))
    
### END SOLUTION

x = np.linspace(-10,10,500)
plot_wave_packet(x,-2)
plot_wave_packet(x,-1)
plot_wave_packet(x,0)
plot_wave_packet(x,1)
plot_wave_packet(x,2)

```

### Part D (5 points)

Now an exercise in two-dimensional ballistics.  We'll start by recalling two textbook formulas for the trajectory $y(x)$ and for the maximum range $R$ of an object moving in 2d under the influence of gravity, starting at $y=0$:

$$
y(x) = x \tan \theta - \frac{gx^2}{2v_0^2 \cos^2 \theta} \\
R = \frac{v_0^2 \sin(2\theta)}{g}
$$

A great way to test these formulas against each other is with a plot!  Let's fix $g=9.8\ m/s^2$ and $v_0 = 6$ m/s, and then __create a plot with the following properties:__

* Show $y(x)$ as a solid line for two trajectories, $\theta = 20$ degrees and $\theta = 45$ degrees.
* Show $R$ for each trajectory as a __dashed, vertical line__ on the same plot, with the _same color_ as the corresponding $y(x)$.
* Include a legend to show which trajectory is which.

The simplest way to add a vertical line to your plot is the `plt.vlines()` function; for things like color and line style, it takes the same options that `plt.plot()` does.

Choosing the plot range, colors, etc. are up to you: the only requirement is that __your plot should be easily usable to see that $y(x)$ and $R$ agree on the maximum range of your projectile__ in both cases.



<!-- grade_id: cell-1d5544cd64a1ce3a, points: 5, grade: true, solution: true -->
```python
import numpy as np
import matplotlib.pyplot as plt

g = 9.8  # m/s^2
v0 = 6.0 # m/s
### BEGIN SOLUTION

x = np.linspace(0,5,50)
def traj(x, v, th):
    return x*np.tan(th) - g*x**2 / (2*v**2 * np.cos(th)**2)
def R(x, v, th):
    return v**2 *np.sin(2*th) / g

theta_20 = 20 * np.pi/180
theta_45 = 40 * np.pi/180

plt.plot(x, traj(x,v0,theta_20), color='blue', label='20 deg')
plt.plot(x, traj(x,v0,theta_45), color='orange', label='40 deg')

plt.vlines(R(x,v0,theta_20),  0, 1, color='blue', linestyle='--')
plt.vlines(R(x,v0,theta_45),  0, 1, color='orange', linestyle='--')

plt.xlim(0,4)
plt.ylim(0,1)

plt.legend()
### END SOLUTION
```