# Tutorial 8: Booleans and Assertions

## PHYS 2600

## T8.1 - Some Boolean logic

Let's do some basic exercises with using Boolean types and the logical operators `and`, `or`, and `not`.  When working in logic, it's often useful to write down a __truth table__, which lists all possible outcomes of a certain logical operation.  Here are the truth tables for `and` and `or`, for your reference as you're working below:

<img src="https://raw.githubusercontent.com/wlough/CU-Phys2600-Fall2025/main/lectures/img/truth-table.png" width=500px />

### Part A

Use Python statements to verify the following more complicated logical statements:

* (T and F) or T --> T
* (F or F) and T --> F
* (F and T) or (F and F) --> F

<!-- grade_id: cell-93f2c8965854b1dd, points: 2, grade: true, solution: true -->
```python
#
```

### Part B

We've seen this flowchart for a recipe for meringue cookies in lecture:

<img src="https://raw.githubusercontent.com/wlough/CU-Phys2600-Fall2025/main/lectures/img/lec1-meringue.png" width=600px />

The original recipe is meant for humans, but it's not hard to imagine trying to implement this for a computer!  Suppose we are writing code for a "smart oven" which will cook this recipe: we're currently working on the step where we decide whether to take the cookies out of the oven or not (the right half of the diagram.)

__Complete the function `finished_baking()` below__, which should return `True` if either condition in the blue box on the right is satisfied (i.e. if oven_min is greater than 45 or firm_to_touch is true), and `False` otherwise.

<!-- grade_id: cell-6491f6674126231d, points: 0, grade: false, solution: true -->
```python
def finished_baking(oven_min, firm_to_touch):
    #
```

This might seem like a silly example to you, but there's no reason this couldn't be real-life working code!  There could be a real timer that sets the variable `oven_min`, and a microcontroller and some sort of probe that sets `firm_to_touch` in a different test routine.

Now in the next cell, __write your own tests__: create three `assert` statements that use the `finished_baking` function and make sure it gives the expected output.

<!-- grade_id: cell-524b8cbda79f2ed5, points: 0, grade: false, solution: true -->
```python
#
```

### Part C

You probably wrote the tests above using the `==` operator, which is normal for a test like this.  But for Boolean data only, __using `==` is almost always redundant.__  For example, suppose we want use an `assert` to make sure that `my_condition` is `True`.  We could write it like this:

```python
my_condition = True
assert my_condition == True
```

__Try running the cell above__ with `my_condition` set to both `True` and `False`; you'll find that it works perfectly.  But so does _this:_

```python
# This code behaves identically to the cell above!

my_condition = True
assert my_condition
```

The statement `my_condition == True` is _identical_ to just `my_condition`, because the equality test maps `True` to `True`, and `False` to `False`!

What if we want to test for a condition being _false_?  Again, we could use `==`, but it turns out that the `not` operator does the same thing:

```python
print(True == False)
print(not True)

print(False == False)
print(not False)
```

So we could just write `assert not thing_that_should_be_false`, and again skip using `==`.  

__Try to get in the habit of skipping `==` when testing Boolean variables!__  It will make your code cleaner, and if you pick good variable names it can be just as clear to read statements like `assert oven_is_preheated` or `assert not oven_is_on_fire`.

__If you used `==` in your `assert` statements above, go back now and rewrite them without `==`.__

### Part D

One fact we mentioned is that Booleans can be interpreted as equivalent to their underlying binary values (0 for `False` and 1 for `True`), and vice-versa.  Usually, you should _not_ use this feature; to avoid surprises from your code, be explicit when carrying out tests to make Booleans!

However, there are exceptions to any rule.  One very useful application of treating Booleans as numbers is _counting_.  We haven't talked about conditional tests in NumPy yet, but if we do a simple test on an array, NumPy will produce an array of Booleans as a result:

```python
## Example code

import numpy as np

a = np.arange(-5, 5)
print(a)
print(a > 0)
```

If you use `np.sum()` on a Boolean array, it will treat the entries as numbers: the result will be a count of how many entries are `True`.  __Try it on the array `a` created in the last cell!__

Then, make another array which contains __all of the _even_ integers from 0 to 10,000__ (remember, `np.arange(a, b, n)` will run from `a` to `b`, stepping by `n`.)  Use the `%` (modulo) operator to count: __how many of these numbers are multiples of 7?__  (You should find that there are 715 of them.)

_(Hint: modulo gives the remainder after dividing.  Any number which is a multiple of 7 will be equal to 0, modulo 7.  For example, `a % 2` will be `0` if a is even, or `1` if a is odd.)_

<!-- grade_id: cell-7074ed75e915a804, points: 0, grade: false, solution: true -->
```python
#
```

Actually, a better function to choose for this application is `np.count_nonzero()`.  This will produce the same result as `np.sum()` on a Boolean array, but it is optimized for this specialized task and so will run more quickly.  Run the cell below to see how they compare, using the Jupyter magic command `%timeit`!

```python
a = np.arange(-100000, 100000)
b = a <= 0

print(np.sum(b))
print(np.count_nonzero(b))

%timeit np.sum(b)
%timeit np.count_nonzero(b)
```

## T8.2 - A cautionary tale about `==`

### Part A

Suppose we want to implement a function `is_triple(x,y)` which returns a Boolean result: `True` if `y` is 3 times `x`, and `False` otherwise.  (Don't ask me _why_ we want such a contrived function, it's just an example.)  Here's one simple way to do it:

```python
def is_triple(x, y):
    return y == 3 * x
```

Looks reasonable, at least for the first few tests, but notice that the fourth test is wrong below:

<!-- grade_id: cell-b0daa4b64035f749, points: 0, grade: true, solution: false -->
```python
## Testing cell - all should be True

print(is_triple(3.0, 9.0))
print(is_triple(1, 3.0))
print(is_triple(1 / 3.0, 1.0))
print(is_triple(1 / 5.0, 3 / 5.0))
```

Why do you think `is_triple` doesn't work as expected in the fourth test case?

<!-- grade_id: cell-a4a75f69a401ea2a, points: 0, grade: true, solution: true -->
__Floating-point precision strikes again!  The order of operations matters at the level of round-off error: 1/5 * 3 gives a slightly different result than 3 / 5 in binary with limited precision.__

### Part B

Implement a better version of this function, `is_approx_triple(x,y,prec)`, where `prec` should start with a default value of `1e-10`.  This function should check to make sure that y is equal to 3x only up to differences of order `prec`.  Then run the testing cell below and make sure everything prints `True`.

<!-- grade_id: cell-d78da5928505435c, points: 0, grade: false, solution: true -->
```python
def is_float_triple(x, y, prec=1e-10):
    #
```

<!-- grade_id: cell-e1e34db5e48ff7fc, points: 0, grade: true, solution: false -->
```python
## Testing cell - all should be True

print(is_float_triple(3.0, 9.0))
print(is_float_triple(1.0, 3.0))
print(is_float_triple(1 / 64.0, 3 / 64.0))
print(is_float_triple(1 / 10.0, 3 / 10.0))
print(is_float_triple(1 / 5.0, 3 / 5.0))
```

Now, __make the default value of `prec` smaller__ until you can reproduce the behavior of the original `is_triple()` function.  Thinking back to [our lecture on floating-point numbers](https://physicscourses.colorado.edu/phys2600/phys2600_sp19/lecture/lec04-numbers-binary/), we recall that the truncation error appears around 1 part in $10^{16}$;  __where do you actually see the behavior change?__

_HINT: there is a common bug that you may have introduced in_ `is_float_triple` - _in fact, I introduced it in my first attempt.  If you're not seeing the expected behavior, try setting `prec=0`.  What should that give you, and what do you see?_

<!-- grade_id: cell-03c258c4f632f7f2, points: 0, grade: false, solution: true -->
```python
is_float_triple(1 / 5.0, 3 / 5.0, prec=1e-16)
```

<!-- grade_id: cell-213d0ede38c08cfa, points: 0, grade: true, solution: true -->
__The average 64-bit truncation error we estimated on tutorial 4 was about $10^{-16}$; that's precisely where you'll see the deviation occur.  Keep in mind that $10^{-16}$ is a relative error - if you try with 1/355. vs. 3/355., you'll see that prec must get down to $10^{-18}$ to cause a False output, since the starting value is around 0.01.__

__Important info:__ for testing, we don't have to do this by hand!  You've probably seen me import the `numpy.testing` module in some test cells already.  There are [a bunch of testing functions available](https://docs.scipy.org/doc/numpy/reference/routines.testing.html) in this module, but the one I commonly use is `numpy.testing.assert_allclose`, which has the following signature:

```
numpy.testing.assert_allclose(actual, desired, rtol=1e-7, atol=0)
```
(and a couple more optional arguments which you will rarely need.)  Here `rtol` is short for "__relative tolerance__", which tests for equality as a fraction of the second argument passed, `desired`.  `atol` is short for "__absolute tolerance__", which is what we did in the example above: it just tests for the difference to be small than `atol`.

To restate it following the documentation: this function carries out the test
```
assert abs(actual - desired) < atol + rtol * abs(desired)
```
Of course, this function has one more advantage over the one we wrote: it works on entire NumPy arrays as well as single numbers.

To avoid things like floating-point truncation error, `rtol` is usually a more appropriate choice!  In general, __only use `atol` if you know to expect some source of absolute error.__

### Part C (optional challenge)

Now that you know how to deal with floating-point numbers in comparisons, let's try a more interesting function!  We can always tell whether a given triangle is a _right triangle_ just by knowing the lengths of the three sides: they will satisfy the Pythagorean theorem $a^2 + b^2 = c^2$ if any angle is 90 degrees.

__Implement the function `is_nearly_right_triangle()` below__, which should return `True` if a triangle with sides `s,t,u` is a right triangle, and `False` otherwise.

_(Hint: remember that you don't know beforehand which of `s,t,u` is the longest side...)_

<!-- grade_id: cell-ebf7ba549e8645e4, points: 0, grade: false, solution: true -->
```python
def is_nearly_right_triangle(s, t, u, prec=1e-6):
    #
```

<!-- grade_id: cell-ad413f4900de05b6, points: 0, grade: true, solution: false -->
```python
import numpy as np

assert not is_nearly_right_triangle(3.0, 3.0, 3.0)
assert is_nearly_right_triangle(3.0, 4.0, 5.0)
assert is_nearly_right_triangle(5.0, 3.0, 4.0)
assert is_nearly_right_triangle(3.0, 5.0, 4.0)
assert is_nearly_right_triangle(np.sqrt(2), 1.0, 1.0)
assert is_nearly_right_triangle(1.4, 1.0, 1.0, prec=0.1)
```