# Tutorial 22: Dictionaries

## PHYS 2600

## T22.1 - Basic use of dictionaries

Let's start by getting some practice using dictionaries and seeing some of their commonly-used methods.

### Part A

__Create a dictionary called `vowel_index`__ that maps the vowels `['a','e','i','o','u']` to their index numbers in the English alphabet (so if we take `a` to correspond to `0`, then we have

| letter | index |
|--------|-------|
| a | 0 |
| e | 4 |
| i | 8 |
| o | 14 |
| u | 20 |





<!-- grade_id: cell-6994cd0401eb349b, points: 2, grade: true, solution: true -->
```python
vowel_index = {
    "a": 0,
    "e": 4,
    "i": 8,
    "o": 14,
    "u": 20,
}
```

As we saw on the cryptography homework, the `ord()` function gives the ASCII index of a character, in which the letters are numbered starting at 97 instead of 0.  __Using `vowel_index`, what is the ASCII index of `'u'`?__  (Once you have an answer, check it with by calling `ord('u')`.)

<!-- grade_id: cell-b2c78f65caa0bfc2, points: 0, grade: false, solution: true -->
```python
print(vowel_index["u"] + 97)
print(ord("u"))
```

Output:
```
117
117
```

The letter `'y'` (index: 24) _sometimes_ acts like a vowel, so let's add it to our dictionary.  __Add `'y'` as a key with value 24 to `vowel_index`__.  Then print out `vowel_index` to see that your code worked.

<!-- grade_id: cell-eea2306839db4f2e, points: 0, grade: false, solution: true -->
```python
vowel_index["y"] = 24
vowel_index
```

Output:
```
{'a': 0, 'e': 4, 'i': 8, 'o': 14, 'u': 20, 'y': 24}
```

### Part B

Any dictionary naturally divides into two sets: the __keys__ (what the index names are), and the __values__ (what objects the index names point to.)  We can access these sets directly with the methods of the same name - __run the cell below__ to see them work.

```python
print(vowel_index.keys())
print(vowel_index.values())
```

Output:
```
dict_keys(['a', 'e', 'i', 'o', 'u', 'y'])
dict_values([0, 4, 8, 14, 20, 24])
```

Note that although these return special dictionary-associated types, we can treat both keys and values like lists.  In particular, we can _iterate_ over them using a `for` loop.  Some practice:

* Use a `for` loop and string concatenation (the `+` operator) to __build a string of all the vowels (and y)__ from the keys of `vowel_index`.
* Use another `for` loop to __calculate the sum of all of the _values_ of `vowel_index`__, i.e. the sum of all vowel letter indices.  (You should get 70.)

<!-- grade_id: cell-e4e7d3c95e8c8c02, points: 0, grade: false, solution: true -->
```python
s = ""
for v in vowel_index.keys():
    s += v

print(s)

total = 0
for i in vowel_index.values():
    total += i

print(total)
```

Output:
```
aeiouy
70
```

__Important note:__  Remember that _order of the keys/values in a dict is not guaranteed!_  In Python, the current implementation of dictionary _happens_ to give keys back in the order they were added to the dict, which results in an alphabetical order.  In the cell below, notice what happens if the keys are entered in a different order.

__Never rely on the order of keys or values in a dictionary - it's meant to be unordered!__  Use a list-sorting function (like `sorted()`) if you need them in a certain order.

```python
vowel_index_2 = {
    "i": 8,
    "a": 0,
    "o": 14,
    "e": 4,
    "u": 20,
}

for v in vowel_index_2.keys():
    print(v)

for v in sorted(vowel_index_2.keys()):
    print("Sorted: ", v)
```

Output:
```
i
a
o
e
u
Sorted:  a
Sorted:  e
Sorted:  i
Sorted:  o
Sorted:  u
```

### Part C

There's a third list which is natural to construct from a dictionary: a list of __key-value pairs__.  To get them, we use the `.items()` method.

Start by __running the cell below__ to see what you get when iterating over `vowel_index.items()`.  Then modify the code to unpack the keys and values separately, and __use `print` to reproduce the block of text__ in the next cell below.

_(Hint: you can either unpack `pair` inside the loop, or use the unpacking notation inside the `for` statement itself: replace the single variable `pair` with a tuple of two variables.)_

<!-- grade_id: cell-e18b932745ea682e, points: 0, grade: false, solution: true -->
```python
for pair in vowel_index.items():
    print(pair)

### BEGIN SOLUTION
for k, v in vowel_index.items():
    print(f"ASCII code for {k}: {v+97}")
### END SOLUTION
```

Output:
```
('a', 0)
('e', 4)
('i', 8)
('o', 14)
('u', 20)
('y', 24)
ASCII code for a: 97
ASCII code for e: 101
ASCII code for i: 105
ASCII code for o: 111
ASCII code for u: 117
ASCII code for y: 121
```

Modify the cell above to print the text below: (if your lines print in a different order, that's fine - dictionaries are unordered!)

```
ASCII code for a: 97
ASCII code for e: 101
ASCII code for i: 105
ASCII code for o: 111
ASCII code for u: 117
ASCII code for y: 121
```

### Part D

In many situations, we might be unsure about whether a key exists in a dictionary.  We can test for its existence using the `in` operator: `x in L` is a Boolean expression which is `True` if `x` is in collection `L`, and `False` otherwise.

The code below makes a new dictionary of Unicode letters, `char_codes`.  Use `in` to answer the following questions:

* Is `ë` in the keys of `char_codes`?
* Is `©` in the values of `char_codes`?
* Is `ȼ` in the keys or values of `char_codes`?


_(Hint: copy and paste these Unicode characters, don't try to figure out how to type them!  You could also look through the dictionary by eye...but it's pretty long!)_

<!-- grade_id: cell-2392f0f5708f9735, points: 0, grade: false, solution: true -->
```python
char_codes = {}
for i in range(512):
    char_codes[chr(i)] = chr(i + 57)

### BEGIN SOLUTION
print("ë" in char_codes.keys())
print("©" in char_codes.values())
print("ȼ" in char_codes.keys())
print("ȼ" in char_codes.values())
### END SOLUTION
```

Output:
```
True
True
False
False
```

Note in passing: the `in` operator automatically checks the keys if we give it a dictionary, so we can just write `c in char_codes` instead of `c in char_codes.keys()`.  (But I had to look up whether this checks the values too - it doesn't - so maybe it's better to avoid ambiguity and just use `.keys()` explicitly!)

### Part E

Back to our simpler dict of vowels, `vowel_index`.  __What happens if we try to get a key that doesn't exist?__  

Ordinarily, this will simply give us an error message, like trying to access a list index that doesn't exist.  __Run the cell below__ to see the message.

```python
vowel_index["z"]
```

Output:
```
KeyError                                  Traceback (most recent call last)
Cell In[9], line 1
----> 1 vowel_index["z"]

KeyError: 'z'
```

We see that we get a `KeyError` above, indicating the key we asked for does not exist.

If we want to use a key that may or may not exist, we can either test with `in` first, but a more streamlined option is to use the `.get()` method to safely ask for that entry.  __Run the cell below__ to see a few variations on the use of `.get()`:

```python
print(vowel_index.get("z"))
print(vowel_index.get("e"))
print(vowel_index.get("z", "default"))
```

Note that by default, `.get()` returns `None` if we try to get a key that doesn't exist.  We can override the default with the second argument to `.get()`.

Below is a list of some letters.  Use a `for` loop to iterate over the list, and use `.get()` with a default argument to __reproduce the text block in the next cell below.__

<!-- grade_id: cell-bfd69f8216bbd335, points: 0, grade: false, solution: true -->
```python
some_letters = ["a", "d", "i", "j", "k", "u"]

### BEGIN SOLUTION
for L in some_letters:
    I = vowel_index.get(L, "not a vowel")
    print(f"{L}: {I}")
### END SOLUTION
```

Add a for loop running through `some_letters`, and use `vowel_index.get()` together with a print statement so that it outputs the following text exactly:
```
a: 0
d: not a vowel
i: 8
j: not a vowel
k: not a vowel
u: 20
```

## T22.2 - Common use patterns for dictionaries

As I noted in lecture, dictionaries are enormously useful for __glue code__ - they enable a lot of basic operations that can help to make your code simpler and more readable.  In the parts below, we'll run through a few example use cases where use of a dictionary is key.


### Part A: Translation

One way to think of a dictionary is as a __map__ from keys to values.  Maps naturally appear in __translations__, such as between different languages; translations match equivalent things to one another.  Here's a table of a few different words in English and French:

| English | French |
|---------|--------|
| car | voiture |
| book | livre |
| coffee | café |
| dog | chien |
| thanks  | merci |

In the cell below, __make a dictionary called `english_to_french`__ with the English words as keys and the French words as values.  You should __build the dictionary with a `for` loop__ using the lists provided.

<!-- grade_id: cell-985daaf6f40841d6, points: 0, grade: true, solution: true -->
```python
english_words = ["car", "book", "coffee", "dog", "thanks"]
french_words = ["voiture", "livre", "café", "chien", "merci"]

### BEGIN SOLUTION
english_to_french = {}
for i in range(5):
    english_to_french[english_words[i]] = french_words[i]
### END SOLUTION
```

<!-- grade_id: cell-32cac0213cb08aea, points: 0, grade: true, solution: false -->
```python
print(english_to_french)
assert english_to_french["car"] == "voiture"
assert english_to_french["thanks"] == "merci"
```

Dictionaries are _one-way_ maps, but a translation is _two-way_: we could also use our map to convert French words back to English, since the pairs are unique.  Fortunately, it's not too difficult to just __invert the dictionary__, i.e. make one that points from values to keys instead.   

__Implement the function `invert_dict(d)` below__, which should return a new dictionary where the _values_ of `d` point to its _keys_.  Then run the following cell to create and test the reverse dictionary `french_to_english`.

_(Hint: start by making a new empty dictionary.  Then use a `for` loop and the `.items()` method to iterate over all key-value pairs `(k,v)` in the original dictionary, and assign them to the empty one with values as keys.)_

<!-- grade_id: cell-1fb27218f2807229, points: 0, grade: false, solution: true -->
```python
def invert_dict(d):
    ### BEGIN SOLUTION
    inv_d = {}
    for k, v in d.items():
        inv_d[v] = k

    return inv_d
    ### END SOLUTION
```

<!-- grade_id: cell-e04c0028cd02064c, points: 0, grade: true, solution: false -->
```python
french_to_english = invert_dict(english_to_french)
print(french_to_english)

assert french_to_english["livre"] == "book"
assert french_to_english["café"] == "coffee"
```

Note: this only makes sense if the map is truly two-way, i.e. all keys and values are matched uniquely __one-to-one__!  If we try to "invert" a dictionary with multiple keys pointing to the same value, we'll lose some information.

If you truly need a bi-directional dictionary for some application in Python (`invert_dict` was simple enough, but if we have a huge data set with millions of entries it will be a massive waste of storage!), you may want to check out [the `bidict` module](https://bidict.readthedocs.io/en/main/).

### Interlude: Joining lists with `zip`

Although doing all of the above with `for` loops is good practice for understanding what is happening, if you want to take two matching lists and join them together into a single dictionary, a much better way is with the built-in keyword `zip`.  `zip` takes two matching lists (or more, but we want two for a dictionary!), and creates a _single_ list containing tuples of pairs.  Here's a quick demonstration:

```python
list_A = [1, 2, 3, 5]
list_B = ["A", "B", "C", "Z"]

list(zip(list_A, list_B))
```

Notice that I have to use `list(zip(...))` since `zip` is another _lazy object_ - we have to force it to evaluate as a list if we want to see the whole thing at once.  But if I instead type-cast it to a `dict()`, then the tuples are interpreted as key-value pairs:

```python
dict(zip(list_A, list_B))
```

In the cell below, __re-create the dict `english_to_french`__ from the lists `english_words` and `french_words` by using `zip()` and `dict()`.  (Much easier than what you did above with `for`!  In fact, if we start with two matched lists, it's easy to make both dictionaries corresponding to mapping in either direction - in other words, it's just as easy to make `french_to_english`, we don't need the `invert_dict` function.)

<!-- grade_id: cell-b7e9992cd38c851a, points: 0, grade: false, solution: true -->
```python
english_to_french = dict(zip(english_words, french_words))
print(english_to_french)
```

### Part B: Descriptive Data

Another common use case for dictionaries is storing __descriptive data__.  This comes in many forms, so it's probably best to just see an example.  Here is a data cell with some information:

```python
# m, F, mu
block_list = (
    (0.5, 0.8, 0.2),
    (1.0, 3.0, 0.2),
    (2.0, 4.4, 0.2),
    (3.0, 5.6, 0.1),
    (2.0, 2.2, 0.3),
    (2.0, 1.6, 0.1),
    (1.0, 2.6, 0.1),
)
block_color_list = ("red", "orange", "yellow", "green", "blue", "white", "black")
```

You may notice that the two lists (tuples, actually) have the same length, which indicates that they are both describing some shared set of data.  The comment indicates that the data in `block_list` are `(m, F, mu)`: short for mass, applied force, and coefficient of static friction.  (In fact, this data describes a simple experimental setup: if we apply the given force to a block with the given mass and static friction, does it move or not?)

Although the comment helps, this is an ideal use case for a dictionary - dictionaries allow us to write __self-documenting code__, in which the data storage medium tells us what is in it!



In the _second cell_ below, I've created a function `check_if_block_moves` that checks if the force of static friction is overcome given the applied force or not.  This function assumes that `blocks` contains a __list of dictionaries__, with each dictionary having keys `m`, `mu`, `F` and `color` to describe a single block.

In the _first_ cell below, __create a list of dictionaries called `blocks`__ so that the cell with `check_if_block_moves()` will run without error (and give the right answer!)  You can do this by hand if you want - it will be tedious! - or use a `for` loop to build the `blocks` list of dictionaries using the data from above.  

_(Hint: Since there are two matching lists, this calls for the __iterator__ pattern.)_  

_(Optional challenge: Actually, there's a very slick way to combine matching lists and iterate over them - we can use the `zip` built-in keyword.  If you're up for a challenge, try to use `zip` to create the `blocks` list of dictionaries without ever writing a list index!)_

<!-- grade_id: cell-dfd5099463b73ddb, points: 0, grade: false, solution: true -->
```python
# By hand:
blocks = [
    {"m": 0.5, "F": 0.8, "mu": 0.2, "color": "red"},
    {"m": 1.0, "F": 3.0, "mu": 0.2, "color": "orange"},
    {"m": 2.0, "F": 4.4, "mu": 0.2, "color": "yellow"},
    {"m": 3.0, "F": 5.6, "mu": 0.1, "color": "green"},
    {"m": 2.0, "F": 2.2, "mu": 0.3, "color": "blue"},
    {"m": 2.0, "F": 1.6, "mu": 0.1, "color": "white"},
    {"m": 1.0, "F": 2.6, "mu": 0.1, "color": "black"},
]

# Or the automated version from the lists above:
blocks = []
for i in range(7):
    b = {}
    b["m"], b["F"], b["mu"] = block_list[i]
    b["color"] = block_color_list[i]
    blocks.append(b)

# And the fancy way using zip:
blocks = []
for (m, F, mu), color in zip(block_list, block_color_list):
    b = {}
    b["m"], b["F"], b["mu"], b["color"] = m, F, mu, color
    blocks.append(b)

print(blocks)
```

```python
# Run this cell after you make the blocks dictionary above!


def check_if_block_moves(block):
    g = 9.8
    F_friction = block["mu"] * block["m"] * g

    if block["F"] <= F_friction:
        print(f"The {block['color']} block doesn't move.")
    else:
        print(f"The {block['color']} block moves!")


for block in blocks:
    check_if_block_moves(block)
```

You should see the following text:

```
The red block doesn't move.
The orange block moves!
The yellow block moves!
The green block moves!
The blue block doesn't move.
The white block doesn't move.
The black block moves!
```

Notice how clean and readable the code in `check_if_block_moves()` is using dictionary-stored descriptive data!  The formula $F_f = \mu m g$, for example, is clearly written out in terms of the data keys `mu` and `m`.

### Part C: Counting _(optional bonus challenge)_

Another common use case for dictionaries is __counting unique values__, for example in creating a histogram of frequencies (how many times something appears in a collection.)

I've provided some data in the array `class_array` below.  For simplicity, each row has four classes in it (real data from real students would have varying amounts of classes, of course.)

Use a dictionary to answer the question: __how many students in our data are enrolled in each class?__  Save your answer in a dictionary called `class_count`, with the class names as keys and the number of students enrolled as values.

(_Hint: this is just a variation on the accumulator design pattern, except that instead of having a single accumulator variable, you use a dictionary.  The value `class_count[k]` will act as an accumulator to count all instances of class `k`._)

```python
import numpy as np

class_array = np.array(
    [
        ["PHYS 2150", "CHEM 1113", "PHYS 2170", "APPM 2350"],
        ["PHYS 2600", "APPM 2350", "PHYS 2150", "PHYS 2170"],
        ["PHYS 2170", "CHEM 1113", "PHYS 2150", "APPM 2350"],
        ["PHYS 2150", "PHYS 2170", "CHEM 1113", "WRTG 3030"],
        ["PHYS 2600", "PHYS 2170", "PHYS 2150", "WRTG 3030"],
        ["CHEM 1113", "PHYS 2150", "PHYS 2170", "WRTG 3030"],
        ["PHYS 2150", "PHYS 2170", "CHEM 1113", "APPM 2350"],
        ["PHYS 2170", "APPM 2350", "PHYS 2150", "CHEM 1113"],
        ["APPM 2350", "PHYS 2150", "PHYS 2600", "WRTG 3030"],
        ["APPM 2350", "PHYS 2170", "PHYS 2600", "PHYS 2150"],
        ["PHYS 2170", "APPM 2350", "PHYS 2150", "CHEM 1113"],
        ["PHYS 2150", "HUEN 3100", "PHYS 2170", "WRTG 3030"],
        ["PHYS 2170", "PHYS 2150", "APPM 2350", "WRTG 3030"],
        ["APPM 2350", "CHEM 1113", "PHYS 2170", "PHYS 2600"],
        ["PHYS 2150", "PHYS 2600", "WRTG 3030", "APPM 2350"],
        ["APPM 2350", "PHYS 2170", "PHYS 2150", "PHYS 2600"],
        ["PHYS 2600", "PHYS 2150", "PHYS 2170", "APPM 2350"],
        ["APPM 2350", "WRTG 3030", "PHYS 2150", "PHYS 2600"],
        ["PHYS 2150", "APPM 2350", "PHYS 2170", "HUEN 3100"],
        ["APPM 2350", "CHEM 1113", "PHYS 2170", "PHYS 2150"],
        ["PHYS 2600", "PHYS 2150", "WRTG 3030", "PHYS 2170"],
        ["APPM 2350", "CHEM 1113", "PHYS 2150", "WRTG 3030"],
        ["PHYS 2170", "PHYS 2150", "CHEM 1113", "PHYS 2600"],
        ["PHYS 2600", "APPM 2350", "PHYS 2170", "PHYS 2150"],
        ["CHEM 1113", "APPM 2350", "PHYS 2150", "PHYS 2170"],
        ["PHYS 2170", "APPM 2350", "CHEM 1113", "HUEN 3100"],
        ["APPM 2350", "WRTG 3030", "PHYS 2600", "HUEN 3100"],
        ["PHYS 2150", "APPM 2350", "PHYS 2170", "WRTG 3030"],
        ["PHYS 2150", "PHYS 2170", "APPM 2350", "PHYS 2600"],
        ["CHEM 1113", "PHYS 2170", "PHYS 2150", "APPM 2350"],
    ]
)
```

<!-- grade_id: cell-bfb37d02cf9e5420, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class_count = {}
for row in class_array:
    for s in row:
        if s in class_count.keys():
            class_count[s] += 1
        else:
            class_count[s] = 1

print(class_count)
### END SOLUTION

keys = set(class_array.ravel())
class_count = {key: 0 for key in keys}
for row in class_array:
    for s in row:
        class_count[s] += 1
print(class_count)
```

Once you've produced the unique count dict `class_count`, __run the cell below__ to make a histogram of the frequencies.  Note that there are a few little plotting tricks here that might be useful to know:

* For data that don't need to be binned, use `plt.bar()` instead of `plt.hist()`.
* To fix overlapping tick labels, use the `rotation=` keyword: here I rotated them to be vertical instead of horizontal.
* Immediately after adding the vertical tick labels, they will be cut off the bottom of the plot.  The command `plt.tight_layout()` is a function which automatically attempts to rescale elements of your plot so that nothing gets cut off.

```python
%matplotlib inline
import matplotlib.pyplot as plt

Nc = len(class_count)
plt.bar(np.arange(Nc), list(class_count.values()))
plt.xticks(np.arange(Nc), class_count.keys(), rotation="vertical")
plt.tight_layout()
```

```python
A = [f"a{_}" for _ in range(7)]
B = [f"b{_}" for _ in range(7)]
C = [f"c{_}" for _ in range(7)]


for i in range(len(A)):
    a = A[i]
    b = B[i]
    print(f"{a=}, {b=}")


for a, b in zip(A, B):
    print(f"{a=}, {b=}")


for i in range(len(A)):
    a = A[i]
    b = B[i]
    c = C[i]
    print(f"{a=}, {b=}, {c=}")


for a, b, c in zip(A, B, C):
    print(f"{a=}, {b=}, {c=}")

```

Output:
```
a='a0', b='b0'
a='a1', b='b1'
a='a2', b='b2'
a='a3', b='b3'
a='a4', b='b4'
a='a5', b='b5'
a='a6', b='b6'
a='a0', b='b0'
a='a1', b='b1'
a='a2', b='b2'
a='a3', b='b3'
a='a4', b='b4'
a='a5', b='b5'
a='a6', b='b6'
a='a0', b='b0', c='c0'
a='a1', b='b1', c='c1'
a='a2', b='b2', c='c2'
a='a3', b='b3', c='c3'
a='a4', b='b4', c='c4'
a='a5', b='b5', c='c5'
a='a6', b='b6', c='c6'
a='a0', b='b0', c='c0'
a='a1', b='b1', c='c1'
a='a2', b='b2', c='c2'
a='a3', b='b3', c='c3'
a='a4', b='b4', c='c4'
a='a5', b='b5', c='c5'
a='a6', b='b6', c='c6'
```